<a href="https://colab.research.google.com/github/srenouf/Tnsi/blob/master/Tnsi_classe_pile_Mini_projet_NPI.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

<!-- grade_id: cell-1702ab9f5c717d14, points: 1, grade: true, solution: false -->
```python
class Pile:
    ''' classe Pile
    création d'une instance Pile avec une liste
    '''
    def __init__(self):
        "Initialisation d'une pile vide"
        self.L=[]

    def vide(self):
        "teste si la pile est vide"
        return self.L==[]

    def empiler(self,x):
        "empile"
        self.L.append(x)

    def depiler(self):
        "dépile"
        assert not self.vide(),"Pile vide"
        return self.L.pop()

    def sommet(self):
        "renvoie le sommet de la pile"
        assert not self.vide(), "Pile vide"
        return self.L[-1]
      
    def indicedernier(self):
        "retourne l'indice du dernier element"
        return(len(self.L))-1
       
    def __str__(self):
        return str(self.L)

p=Pile()
p.empiler(5)
print(p)
```

Output:
```
[5]
```